### Arithmetic Expressions with Parsing Expression Grammar

Consider the following PEG for arithmetic expressions:

    S ← E !.
    E ← T ('+' T)*
    T ← F ('×' F)*
    F ← P'²' / P
    P ← 'a' / '(' E ')'

Implement it as a backtracking parser, following the rules from the course notes. You may use the commented version of method `literal`.

<!-- grade_id: cell-d26520c7eddbab9b, points: 3, grade: true, solution: true -->
```python
class Backtrack:
    src: str

    def literal(self, k: int, a: str):
        if self.src.startswith(a, k): return k + len(a)
        # if self.src.startswith(a, k): print('match', a, 'at', k); return k + len(a)
        # else: print('fail', a, 'at', k); return None

    def S(self, k: int):
        k = self.E(k)
        r = k; k = self.literal(k, '.')
        k = r if k == None else None
        if k != None and k < len(self.src):
            k = None
        return k

    def E(self, k: int):
        k = self.T(k)
        r = k
        while k != None:
            r = k
            k = self.literal(k, '+')
            if k != None: k = self.T(k)
        k = r
        return k

    def T(self, k: int):
        k = self.F(k)
        r = k
        while k != None:
            r = k
            k = self.literal(k, '×')
            if k != None: k = self.F(k)
        k = r
        return k
        
    def F(self, k: int):
        r = k
        k = self.P(k)
        k = self.literal(k, '²')
        if k == None:
            k = r
            k = self.P(k)
        return k

    def P(self, k: int):
        r = k
        k = self.literal(k, 'a')
        if k == None:
            k = r
            k = self.literal(k, '(')
            if k != None: k = self.E(k)
            if k != None: k = self.literal(k, ')')
        return k
    
    def parse(self, s: str):
        self.src = s; return self.S(0)
        
```

```python
p = Backtrack()
```

<!-- grade_id: cell-2defb949e3f95c23, points: 1, grade: true, solution: false -->
```python
assert p.parse('a²') == 2
assert p.parse('a²+a²') == 5
assert p.parse('a²+a+a²') == 7
```

<!-- grade_id: cell-100652ba17ebb90d, points: 1, grade: true, solution: false -->
```python
assert p.parse('((a)²)') == 6
assert p.parse('(a²+a)×a') == 8
assert p.parse('a×a²×a') == 6
```

<!-- grade_id: cell-35c67e76c99ad61e, points: 1, grade: true, solution: false -->
```python
assert p.parse('a²²') == None
assert p.parse('a+²a') == None
assert p.parse('a(+)a') == None
```