### Grammar for `aⁱbⁿ` (NLTK)

As a reminder, grammar `G = (T, N, P, S)` with `T = {a, b}`, `N = {S}`, and productions `P`

    S → ε
    S → aSbS
    S → bSaS

is written in NLTK with productions for the same nonterminal on one line with `|` or separate lines; `ε` is simply left out:

```python
import nltk
grammar = nltk.CFG.fromstring("""
S -> 'a' S 'b' S | 'b' S 'a' S | 
""")
parser = nltk.ChartParser(grammar)
trees = list(parser.parse(['a', 'b', 'a', 'b']))
for t in trees: print(t)
```

The start symbol is the left-hand side of the first production. The function `parser.parse(sent)` returns a generator for parse trees, which above is used to produce a list of parse trees. If `sent` cannot be parsed, the list is empty. The parameter `['a', 'b', 'a', 'b']` can be abbreviated to `'abab'`:

```python
assert list(parser.parse('abab')) != []
assert list(parser.parse('aba')) == []
```

---

Using NLTK, write a grammar for the language `aⁱbⁿ` where `0 ≤ i < n`!

<!-- grade_id: cell-1a049d553b84f97b, points: 0, grade: false, solution: true -->
```python
import nltk

grammar = nltk.CFG.fromstring("""
S -> 'b' | 'a' S 'b' | S 'b' 
""")

# S -> b | aSb | Sb

# aSb
# a(aSb)b
# aabbb

# aSb
# a(Sb)b
# ab
# b
# a(aSbb)bb
# abb
# aabbb
# abbb
# aaaabbbbb

# parser = nltk.ChartParser(grammar)
# trees = list(parser.parse(['a', 'b', 'a', 'b']))
# for t in trees: print(t)
parser = nltk.ChartParser(grammar)
```

You can insert cells to see the trees NLTK generates; the grammar can be ambiguous. Here are some tests:

<!-- grade_id: cell-442835edb917c0dd, points: 2, grade: true, solution: false -->
```python
assert list(parser.parse('')) == []
assert list(parser.parse('a')) == []
assert list(parser.parse('b')) != []
assert list(parser.parse('aa')) == []
assert list(parser.parse('ab')) == []
assert list(parser.parse('ba')) == []
assert list(parser.parse('bb')) != []
```

<!-- grade_id: cell-4c38ec0a90f280e9, points: 2, grade: true, solution: false -->
```python
assert list(parser.parse('aaa')) == []
assert list(parser.parse('aab')) == []
assert list(parser.parse('abb')) != []
assert list(parser.parse('bbb')) != []
```

```python

```