#### Lowercasing HTML Tags (sed)

Consider `IMG` tags in HTML files, as in:

```python
%%writefile data/q.html
<img src="PiCtuRe.PnG "/>
<img src="PiCtuRe.PnG"></img>
<img src="PiCtuRe.PnG">alt</img>
<img src="PiCtuRe.PnG"> alt   </img>
<img src ="PiCtuRe.PnG" />
<img src = "PiCtuRe.PnG"/>
<img onclick="alert('Clicked!')" src = "PiCtuRe.PnG"/>
```

Output:
```
Overwriting data/q.html
```

 Write an `sed` command to lowercase `SRC` values. For example, `<img   src="PiCtuRe.PnG">` should be replaced by `<img   src="picture.png">`. Note that there can be arbitrary space around `src` and `=` and before `>`. [2 points]

<!-- grade_id: cell-340fa4897e6a08b6, points: 0, grade: false, solution: true -->
```python
%%capture output
%%bash
#sed -r 's/TODO/TODO/g' data/q.html
sed -rn 's/(<img.*src\s*=\s*)(".*")(\s*\/?>.*)/\1\L\2\3/p' data/q.html
```

```python
print(output) # for testing
```

Output:
```
<img src="picture.png "/>
<img src="picture.png"></img>
<img src="picture.png">alt</img>
<img src="picture.png"> alt   </img>
<img src ="picture.png" />
<img src = "picture.png"/>
<img onclick="alert('Clicked!')" src = "picture.png"/>
```

<!-- grade_id: cell-b749bd831640221b, points: 2, grade: true, solution: false -->
```python
assert str(output) == """<img src="picture.png "/>
<img src="picture.png"></img>
<img src="picture.png">alt</img>
<img src="picture.png"> alt   </img>
<img src ="picture.png" />
<img src = "picture.png"/>
<img onclick="alert('Clicked!')" src = "picture.png"/>
"""
```